<problem display_name="Energy Levels">
<script type="loncapa/python">
<![CDATA[
import vectordraw

### Custom Checks ###

def check_vector_count(check, vectors):
    if len(vectors) != check['expected']:
        return 'Number of vectors is not correct.'

def check_vector_angles(check, vectors):
    for vec_name in vectors:
        angle_check = {
            'vector': vec_name,
            'expected': check['expected'],
            'tolerance': check['tolerance']
        }
        result = check_angle(angle_check, vectors)
        if result is not None:
            return result

def check_energy_vector(check, vectors):
    vector_found = False
    for vec_name in vectors:
        if check_tail_y({'vector': vec_name, 'expected': check['tail_y']}, vectors):
            continue
        elif check_length({'vector': vec_name, 'expected': check['length']}, vectors):
            continue
        else:
            vector_found = True
            break
    if not vector_found:
        return 'Some of the transition vectors are not correct.'


vectordraw.custom_checks = {
    'vector_count': check_vector_count,
    'vector_angles': check_vector_angles,
    'energy_vector': check_energy_vector
}
vectordraw.success_message = 'Good job!'
]]>
</script>
<p>
When an electron moves from the n=1 level to the n=3 level, there are multiple paths for the electron to return to the ground state. In the activity above, draw vectors indicating these paths. 
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/4_energyLevels.html"
    />
</customresponse>
</problem>
