<problem display_name="Force Diagram - String and Spring">

<script type="loncapa/python">
<![CDATA[
import vectordraw
grader = vectordraw.Grader()
]]>
</script>
<p><img src="/static/Davidson_EdX-50.png" width="901" height="112"/></p>
<p>
A student attaches two fans on a 0.700 kg cart that is resting on a cart track. The fans are pointed in opposite directions as shown in the figure below. When each fan is running, it pushes the air in an outward direction from the center of the cart. The student turns on both fans, each of which produces an external force of 1.0 N on the cart.
Let \(F_G\) represents the gravitational force, \(F_N\) the normal force, and \(F_1\) and \(F_2\) the external forces exerted on the cart when Fans 1 and 2 are running. </p>
<p>
  Draw the appropriate Free Body Diagram representing the described system.
</p>
<img src="/static/IMG_4.2.2_Two_Fan_FBDs_v2.jpg" />
<p>
  <i>This is a vector drawing problem type. Use your mouse, or the "add vector" button, to place and drag vectors to the blue points. Be sure to place the correctly labeled vector at the appropriate location. </i>
</p>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/11_FDiagrams_LC2_Q2.html"
    />
</customresponse>
  <solution>
<div class="detailed-solution">
<p>Explanation</p>
  <img src="/static/IMG_4.1.2_1_1.jpg" width="400" height="300"/>
<p>[explanation]
Fan 1 blows outward from the cart, so the air would create an force on the cart directed to the right. Fan 2 also blows outward, but the force it creates would be to the left. All lengths should be relative to the amount to force calculations.
</p>

</div>
</solution>
</problem>
