<problem display_name="Electric Field">
<script type="loncapa/python">
<![CDATA[
import vectordraw
]]>
</script>
<p>
The below diagram consists of two parallel charge distributions. Use your mouse to draw vectors indicating the direction of the electric field at points A, B, C, D, and E.
</p>
<p>
  <i>This is a vector drawing problem type. Use your mouse, or the "add vector" button, to place vectors at the blue points. Be sure to place the correctly labeled vector at the appropriate location. </i>
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/9_capacitorField.html"
    />
</customresponse>
</problem>
