<problem display_name="Draw Force Diagram">
<script type="loncapa/python">
<![CDATA[

import vectordraw

### Custom Checks ###

def check_y_intercept(check, vectors):
    vec = vectors[check['vector']]
    tolerance = check.get('tolerance', 1.0)
    expected = check['expected']
    slope = (vec.tip.y-vec.tail.y)/(vec.tip.x-vec.tail.x)
    y_intercept = vec.tail.y - slope*vec.tail.x
    dist = abs(expected - y_intercept)
    if dist > tolerance:
        return "Hmm, is your tangent line touching the curve?"
	
custom_checks = {
	'y_intercept': check_y_intercept,
}
success_message = 'Good job!'

grader = vectordraw.Grader(success_message, custom_checks)
]]>
</script>

<p>Using the slope and intercept that you found above, draw the line tangent to the curve in the first quadrant at \(x=1\). Does your answer make sense?</p>

<p><em>Stretch and orient your tangent line by dragging its ends; shift your tangent line by dragging the line's center. em</em></p>


<customresponse cfn="vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="480"
    html_file="/static/ID_draw_LC1_HOA1.html"
    />
</customresponse>
</problem>

