<problem display_name="Vector Draw - Protractor">
<script type="loncapa/python">
<![CDATA[
import vectordraw
]]>
</script>
<p>
Draw a vector whose angle is 31 degrees relative to the positive x direction. <i>Note,
be sure that your vector starts from the center of the protractor, and ends at the outsideedge with the appropriate angle.</i>
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="480"
    html_file="/static/1_protractor.html"
    />
</customresponse>
</problem>

