<problem display_name="Gravitational Fields">
<script type="loncapa/python">
<![CDATA[
import vectordraw
]]>
</script>
<p>
Draw the direction of the gravitational field at points A and B in the diagram below.  
</p>
<p>
<i>This is a vector drawing problem type. Use your mouse, or the "add vector" button, to place vectors at the blue points. Be sure to place the correctly labeled vector at the appropriate location. </i>
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/9_gravTreeSimple.html"
    />
</customresponse>
</problem>
