<problem display_name="Electric Field">
<script type="loncapa/python">
<![CDATA[
import vectordraw
grader = vectordraw.Grader()
]]>
</script>
<p>
Below is a diagram containing equipotential lines for an unknown charge configuration. Use your mouse to draw vectors indicating the direction of the electric field at points A, B, and C.
</p>
<p>
  <i>This is a vector drawing problem type. Use your mouse, or the "add vector" button, to place vectors at the blue points. Be sure to place the correctly labeled vector at the appropriate location. </i>
</p>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/6_pointCharge.html"
    />
</customresponse>
</problem>
