<problem display_name="Force Diagram - String and Spring">

<script type="loncapa/python">
<![CDATA[
import vectordraw

def check_compare_normal_and_gravity_lengths(check, vectors):
    normal = vectors['N']
    gravity = vectors['g']
    if (gravity.length - normal.length) < 0.5:
        return 'Gravitational Force should be larger than the Normal Force.'

vectordraw.custom_checks = {
  'compare_normal_and_gravity_lengths': check_compare_normal_and_gravity_lengths
}

vectordraw.success_message = 'Good job!'
]]>
</script>
<img src="/static/Davidson_EdX-51.png" width="901" height="112"/>
<p>Let \(F_G\) represent the gravitational force, and\(F_N\) represent the normal force. Draw the free-body diagram <strong>using only \(F_G\) and \(F_N\)</strong>.</p>
<p>
  <i>This is a vector drawing problem type. Use your mouse, or the "add vector" button, to place and drag vectors to the blue points. Be sure to place the correctly labeled vector at the appropriate location. </i>
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/12_FDiagrams_LC3_Q2.html"
    />
</customresponse>
  <solution>
<div class="detailed-solution">
<p>Explanation</p>
<p>[explanation]
The gravitational force, \(F_G\), should point straight downward. The normal force,\(F_N\), should be perpendicular to the ramp. However, the normal force should be of smaller magnitude than the gravitational force because the ramp is at an angle. A component of the gravitational force is directed parallel to the ramp, leaving less of the gravitational force to push the cart into the ramp. Therefore, the ramp needs to push upward on the cart with less force, and the normal force is reduced relative to what it would be if the ramp were flat.
</p>

</div>
</solution>
</problem>
