<problem display_name="Flight Path">
<script type="loncapa/python">
<![CDATA[
import vectordraw
grader = vectordraw.Grader()
]]>
</script>
<p>
The image below shows the flight path of a procectile. At points A, B, and C, draw the appropriatly labeled velocity vector. 
</p>
<p>
  <i>Be sure that the "tail" of each vector starts as close as possible to points A, B, and C.</i>
</p>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="750"
    height="480"
    html_file="/static/3_flightPath.html"
    />
</customresponse>
</problem>
