<problem display_name="webGLDemo">
<script type="loncapa/python">
<![CDATA[
import vectordraw

def check_if_correct_vectors_plotted(check, vectors):
    suppply = vectors['S']
    if 'D' in vectors or 'LRAS' in vectors:
        return 'You have plotted more than the requested number of curves.'

custom_checks = {
  'check_if_correct_vectors_plotted': check_if_correct_vectors_plotted
}

success_message = 'Good job!'

grader = vectordraw.Grader(success_message, custom_checks)
]]>
</script>
<p>
Draw your typical supply / demand relationship.
</p>
<p>
  <i>Testing segment tool.</i>
</p>
<customresponse cfn="grader.cfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="480"
    html_file="/static/segmentMacro.html"
    />
</customresponse>
</problem>
