<problem display_name="Forcd Diagram - Incline">
<script type="loncapa/python">
<![CDATA[
import vectordraw
]]>
</script>
<p>
A box remains stationary on an inclined plane at an angle of 20 degrees with respect to the x-axis. Draw all appropriate vectors that define the force diagram for this sytem.
</p>
<p>
  <i>Be sure that the "tail" of each vector starts at the center of mass of the box.</i>
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="500"
    html_file="/static/2_boxIncline.html"
    />
</customresponse>
</problem>
