<problem display_name="webGLDemo">
<script type="loncapa/python">
<![CDATA[
import vectordraw

### Custom Checks ###

def check_sras2_midpoint(check, vectors):
    vec = vectors['SRAS2']
    tolerance = check.get('tolerance', 1.0)
    mid_x = (vec.tail.x + vec.tip.x) / 2.0
    mid_y = (vec.tail.y + vec.tip.y) / 2.0
    expected_x = check.get('x', mid_x)
    expected_y = check.get('y', mid_y)
    dist = math.hypot(expected_x - mid_x, expected_y - mid_y)
    if dist > tolerance:
        return "You have drawn the curve incorrectly."

vectordraw.custom_checks = {
  'check_sras2_midpoint': check_sras2_midpoint
}


vectordraw.success_message = 'Good job!'


]]>
</script>
<p>
<b>Scenario 3: Now assume that in Econoland input costs decline. Using an upward sloping short-run aggregate supply curve, graph this change.</b>
</p>
<p>
Use your mouse (or finger on a tablet), to draw the new SRAS curve.
</p>
<customresponse cfn="vectordraw.vglcfn">
  <jsinput gradefn="getInput"
    get_statefn="getState"
    set_statefn="setState"
    initial_state='{}'
    width="800"
    height="480"
    html_file="/static/drawSRAS_S3.html"
    />
</customresponse>
</problem>
